# Getting started

Topics:
1.  How to print
2. How to assign variables
3. How to comment
4. Syntax for basic math

## How to print

In Julia we usually use `println()` to print.

```julia
println("I'm excited to learn Julia!")
```

## How to assign variables

```julia
my_answer = 42
```

```julia
typeof(my_answer)
```

```julia
my_pi = 3.14159
typeof(my_pi)
```

```julia
😺 = "smiley cat!"
typeof(😺)
```

To type a smiley cat, use tab completion to select the emoji name and then tab again

```julia
# \:smi + <tab> --> select with down arrow + <enter> ---> <tab> + <enter> to complete
```

After assigning a value to a variable, we can reassign a value of a different type to that variable without any issue.

```julia
😺 = 1
```

```julia
typeof(😺)
```

The same entry method can be used for e.g. Greek letters and other symbols.

```julia
# \alpha<TAB>
α = 1
```

```julia
# \approx<TAB>
α ≈ 1.0000000000000002
```

## Multiple assignment

```julia
a = b = 0
```

```julia
a, b = 2, 3
```

```julia
a, b = b, a
```

## How to comment

```julia
# You can leave comments on a single line using the pound/hash key
```

```julia
#=

For multi-line comments, 
use the '#= =#' sequence.

=#
```

## Syntax for basic math

```julia
sum = 3 + 7
```

```julia
difference = 10 - 3
```

```julia
product = 20 * 5
```

You can multiply implicitly with numeric literals:

```julia
x = 3
2x
```

```julia
quotient = 100 / 10
```

```julia
power = 10^2
```

```julia
modulus = 13 % 5
```

## Rational numbers

```julia
1//2 + 2//3
```

## Complex numbers

```julia
(1 + 2im)^2
```

## High precision

```julia
big(2)^200
```

```julia
big(pi)
```

```julia
big"0.1"
```

```julia
big(0.1)
```

## Other numeric types

```julia
typeof(0xff)
```

```julia
typeof(0x10ab)
```

```julia
typeof(1.0)
```

```julia
typeof(1.0f0)
```

You can use underscore as a digit separator:

```julia
1_000_000
```

## Equality

`==` is "conventional" equality:

```julia
1 == 1.0
```

`===` determines whether two objects are *indistinguishable*:

```julia
1 === 1.0
```

```julia
2 === 1+1
```

```julia
"Hello, world" === join(["Hello", "world"], ", ")
```

### Exercises

#### 1.1 
Look up docs for the `convert` function.

```julia

```

#### 1.2 
Assign `365` to a variable named `days`. Convert `days` to a float and assign it to variable `days_float`

```julia

```

<!-- grade_id: cell-715f78016beb0489, points: 1, grade: true, solution: false -->
```julia
@assert days == 365
@assert days_float == 365.0

```

#### 1.3 
See what happens when you execute

```julia
convert(Int64, "1")
```
and

```julia
parse(Int64, "1")
```

```julia
Please click on `Validate` on the top, once you are done with the exercises.
```