# Getting started

Topics:
1.  How to print
2. How to assign variables
3. How to comment
4. Syntax for basic math

## How to print

In Julia we usually use `println()` to print.

```julia
println("I'm excited to learn Julia!")
```

Output:
```
I'm excited to learn Julia!
```

## How to assign variables

```julia
my_answer = 42
```

Output:
```
42
```

```julia
typeof(my_answer)
```

Output:
```
Int64
```

```julia
my_pi = 3.14159
typeof(my_pi)
```

Output:
```
Float64
```

```julia
😺 = "smiley cat!"
typeof(😺)
```

Output:
```
String
```

To type a smiley cat, use tab completion to select the emoji name and then tab again

```julia
# \:smi + <tab> --> select with down arrow + <enter> ---> <tab> + <enter> to complete
```

After assigning a value to a variable, we can reassign a value of a different type to that variable without any issue.

```julia
😺 = 1
```

Output:
```
1
```

```julia
typeof(😺)
```

Output:
```
Int64
```

The same entry method can be used for e.g. Greek letters and other symbols.

```julia
# \alpha<TAB>
α = 1
```

Output:
```
1
```

```julia
# \approx<TAB>
α ≈ 1.0000000000000002
```

Output:
```
true
```

```julia
?isapprox
```

Output:
```
search: [0m[1mi[22m[0m[1ms[22m[0m[1ma[22m[0m[1mp[22m[0m[1mp[22m[0m[1mr[22m[0m[1mo[22m[0m[1mx[22m

[36m  isapprox(x, y; rtol::Real=atol>0 ? 0 : √eps, atol::Real=0, nans::Bool=false, norm::Function)[39m

  Inexact equality comparison: [36mtrue[39m if [36mnorm(x-y) <= max(atol,
  rtol*max(norm(x), norm(y)))[39m. The default [36matol[39m is zero and the default [36mrtol[39m
  depends on the types of [36mx[39m and [36my[39m. The keyword argument [36mnans[39m determines
  whether or not NaN values are considered equal (defaults to false).

  For real or complex floating-point values, if an [36matol > 0[39m is not specified,
  [36mrtol[39m defaults to the square root of [36meps[39m of the type of [36mx[39m or [36my[39m, whichever is
  bigger (least precise). This corresponds to requiring equality of about half
  of the significand digits. Otherwise, e.g. for integer arguments or if an
  [36matol > 0[39m is supplied, [36mrtol[39m defaults t
```

## Multiple assignment

```julia
a = b = 0
```

Output:
```
0
```

```julia
a, b = 2, 3
```

Output:
```
(2, 3)
```

```julia
a, b = b, a
```

Output:
```
(3, 2)
```

```julia
a
```

Output:
```
3
```

## How to comment

```julia
# You can leave comments on a single line using the pound/hash key
```

```julia
#=

For multi-line comments, 
use the '#= =#' sequence.

=#
```

## Syntax for basic math

```julia
sum = 3 + 7
```

Output:
```
10
```

```julia
difference = 10 - 3
```

Output:
```
7
```

```julia
product = 20 * 5
```

Output:
```
100
```

You can multiply implicitly with numeric literals:

```julia
x = 3
2x
```

Output:
```
6
```

```julia
quotient = 100 ÷ 11
```

Output:
```
9
```

```julia
power = 10^2
```

Output:
```
100
```

```julia
modulus = 13 % 5
```

Output:
```
3
```

## Rational numbers

```julia
1//2 + 2//3
```

Output:
```
7//6
```

## Complex numbers

```julia
(1 + 2im)^2
```

Output:
```
-3 + 4im
```

## High precision

```julia
big(2)^2000
```

Output:
```
114813069527425452423283320117768198402231770208869520047764273682576626139237031385665948631650626991844596463898746277344711896086305533142593135616665318539129989145312280000688779148240044871428926990063486244781615463646388363947317026040466353970904996558162398808944629605623311649536164221970332681344168908984458505602379484807914058900934776500429002716706625830522008132236281291761267883317206598995396418127021779858404042159853183251540889433902091920554957783589672039160081957216630582755380425583726015528348786419432054508915275783882625175435528800822842770817965453762184851149029376
```

```julia
big(pi)
```

Output:
```
3.141592653589793238462643383279502884197169399375105820974944592307816406286198
```

```julia
big"0.1"
```

Output:
```
1.000000000000000000000000000000000000000000000000000000000000000000000000000002e-01
```

```julia
big(0.1)
```

Output:
```
1.000000000000000055511151231257827021181583404541015625e-01
```

## Other numeric types

```julia
0xff::UInt8
```

Output:
```
0xff
```

```julia
typeof(0x10ab)
```

Output:
```
UInt16
```

```julia
typeof(1.0)
```

Output:
```
Float64
```

```julia
typeof(1.0f0)
```

Output:
```
Float32
```

You can use underscore as a digit separator:

```julia
1_000_000
```

Output:
```
1000000
```

## Equality

`==` is "conventional" equality:

```julia
1 == 1.0
```

Output:
```
true
```

`===` determines whether two objects are *indistinguishable*:

```julia
1 === 1.0
```

Output:
```
false
```

```julia
2 === 1+1
```

Output:
```
true
```

```julia
"Hello, world" === join(["Hello", "world"], ", ")
```

Output:
```
true
```

### Exercises

#### 1.1 
Look up docs for the `convert` function.

```julia
?convert
```

Output:
```
search: [0m[1mc[22m[0m[1mo[22m[0m[1mn[22m[0m[1mv[22m[0m[1me[22m[0m[1mr[22m[0m[1mt[22m [0m[1mc[22m[0m[1mo[22mde_[0m[1mn[22mati[0m[1mv[22m[0m[1me[22m @[0m[1mc[22m[0m[1mo[22mde_[0m[1mn[22mati[0m[1mv[22m[0m[1me[22m

[36m  convert(T, x)[39m

  Convert [36mx[39m to a value of type [36mT[39m.

  If [36mT[39m is an [36mInteger[39m type, an [36mInexactError[39m will be raised if [36mx[39m is not
  representable by [36mT[39m, for example if [36mx[39m is not integer-valued, or is outside
  the range supported by [36mT[39m.

[1m  Examples[22m
[1m  ≡≡≡≡≡≡≡≡≡≡[22m

[36m  julia> convert(Int, 3.0)[39m
[36m  3[39m
[36m  [39m
[36m  julia> convert(Int, 3.5)[39m
[36m  ERROR: InexactError: Int64(Int64, 3.5)[39m
[36m  Stacktrace:[39m
[36m  [...][39m

  If [36mT[39m is a [36mAbstractFloat[39m or [36mRational[39m type, then it will return the closest
  value to [36mx[39m representable by [36mT[39m.

[36m  julia
```

#### 1.2 
Assign `365` to a variable named `days`. Convert `days` to a float and assign it to variable `days_float`

```julia
days = 365
```

Output:
```
365
```

```julia
days_float = convert(Float64, days)
```

Output:
```
365.0
```

```julia

```

```julia

```

<!-- grade_id: cell-715f78016beb0489, points: 1, grade: true, solution: false -->
```julia
@assert days == 365
@assert days_float == 365.0

```

#### 1.3 
See what happens when you execute

```julia
convert(Int64, "1")
```
and

```julia
parse(Int64, "1")
```

```julia
Please click on `Validate` on the top, once you are done with the exercises.
```

```julia
convert(Int64, "1")
```

Output:
```
Stacktrace:
 [1] top-level scope at In[56]:1
```

```julia
parse(Int64, "1")
```

Output:
```
1
```

```julia

```