# Strings

Topics:
1. How to get a string
2. String interpolation
3. String concatenation

## How to get a string

Enclose your characters in " " or """ """!

```julia
s1 = "I am a string."
```

Output:
```
"I am a string."
```

```julia
s2 = """I am also a string. """
```

Output:
```
"I am also a string. "
```

There are a couple functional differences between strings enclosed in single and triple quotes. <br>
One difference is that, in the latter case, you can use quotation marks within your string.

```julia
"Here, we get an "error" because it's ambiguous where this string ends "
```

Output:
```
syntax: cannot juxtapose string literal
```

```julia
"""Look, Mom, no "errors"!!!"""
```

Output:
```
"Look, Mom, no \"errors\"!!!"
```

Note that ' ' define a character, but NOT a string!

```julia
typeof('a')
```

Output:
```
Char
```

```julia
'We will get an error here'
```

Output:
```
syntax: invalid character literal
```

## String interpolation

We can use the $ sign to insert existing variables into a string and to evaluate expressions within a string. <br>
Below is an example that contains some highly sensitive personal information.

```julia
name = "Jane"
num_fingers = 10
num_toes = 10
```

Output:
```
10
```

```julia
println("Hello, my name is $name.")
println("I have $num_fingers fingers and $num_toes toes.")
```

Output:
```
Hello, my name is Jane.
I have 10 fingers and 10 toes.
```

```julia
println("That is $(num_fingers + num_toes) digits in all!!")
```

Output:
```
That is 20 digits in all!!
```

## String concatenation

Below are three ways we can concatenate strings! <br><br>
The first way is to use the `string()` function. <br>
`string()` converts non-string inputs to strings.

```julia
s3 = "How many cats ";
s4 = "is too many cats?";
😺 = 10
```

Output:
```
10
```

```julia
string(s3, s4)
```

Output:
```
"How many cats is too many cats?"
```

```julia
string("I don't know, but ", 😺, " is too few.")
```

Output:
```
"I don't know, but 10 is too few."
```

We can also use `*` for concatenation!

```julia
s3*s4
```

Output:
```
"How many cats is too many cats?"
```

### Exercises

#### 2.1 
Create a string that says "hi" 1000 times, first with `repeat` and then with the exponentiation operator, which can call `*` under the hood. Assign it the variable `hi` below.

<!-- grade_id: cell-5b38bdb1115161b4, points: 0, grade: false, solution: true -->
```julia
?repeat
```

Output:
```
search: [0m[1mr[22m[0m[1me[22m[0m[1mp[22m[0m[1me[22m[0m[1ma[22m[0m[1mt[22m

[36m  repeat(A::AbstractArray, counts::Integer...)[39m

  Construct an array by repeating array [36mA[39m a given number of times in each
  dimension, specified by [36mcounts[39m.

[1m  Examples[22m
[1m  ≡≡≡≡≡≡≡≡≡≡[22m

[36m  julia> repeat([1, 2, 3], 2)[39m
[36m  6-element Array{Int64,1}:[39m
[36m   1[39m
[36m   2[39m
[36m   3[39m
[36m   1[39m
[36m   2[39m
[36m   3[39m
[36m  [39m
[36m  julia> repeat([1, 2, 3], 2, 3)[39m
[36m  6×3 Array{Int64,2}:[39m
[36m   1  1  1[39m
[36m   2  2  2[39m
[36m   3  3  3[39m
[36m   1  1  1[39m
[36m   2  2  2[39m
[36m   3  3  3[39m

  ────────────────────────────────────────────────────────────────────────────

[36m  repeat(A::AbstractArray; inner=ntuple(x->1, ndims(A)), outer=ntuple(x->1, ndims(A)))[39m

  Construct an array by repeating the entries of [36mA[39m. The i-th element of [36minner[39m
  specifies the n
```

```julia
"hi"^1000
```

Output:
```
"hihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihih
```

<!-- grade_id: cell-20b1027a968fbb48, points: 1, grade: true, solution: false -->
```julia
@assert hi == "hihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihihi"
```

#### 2.2 
Declare two variables

```julia
a = 3
b = 4
```
and use them to create two strings:
```julia
"3 + 4"
"7" 
```
and store the results in `c` and `d` respectively

<!-- grade_id: cell-36beb524c8ecd33b, points: 0, grade: false, solution: true -->
```julia
a=3
b=4
"$a + $b"
```

Output:
```
"3 + 4"
```

```julia
string(a+b)
```

Output:
```
"7"
```

<!-- grade_id: cell-14c3c5b015a97892, points: 1, grade: true, solution: false -->
```julia
@assert c == "3+4"
@assert d == "7"
```

Please click on `Validate` on the top, once you are done with the exercises.