# Functions

Topics:
1. How to declare a function
2. Duck-typing in Julia
3. Mutating vs. non-mutating functions
4. Some higher order functions

## How to declare a function
Julia gives us a few different ways to write a function. The first requires the `function` and `end` keywords

```julia
function sayhi(name)
    println("Hi $name, it's great to see you!")
end
```

Output:
```
sayhi (generic function with 1 method)
```

```julia
function f(x)
    x^2 + 2x .+ 1
end
```

Output:
```
f (generic function with 1 method)
```

We can call either of these functions like this:

```julia
sayhi("C-3PO")
```

Output:
```
Hi C-3PO, it's great to see you!
```

```julia
f(42)
```

Output:
```
1849
```

Alternatively, we could have declared either of these functions in a single line

```julia
sayhi2(name) = println("Hi $name, it's great to see you!")
```

Output:
```
sayhi2 (generic function with 1 method)
```

```julia
f2(x) = x^2 + 2x + 1
```

Output:
```
f2 (generic function with 1 method)
```

```julia
sayhi2("R2D2")
```

Output:
```
Hi R2D2, it's great to see you!
```

```julia
f2(42)
```

Output:
```
1849
```

Finally, we could have declared these as "anonymous" functions

```julia
sayhi3 = name -> println("Hi $name, it's great to see you!")
```

Output:
```
#3 (generic function with 1 method)
```

```julia
f3 = x -> x^2 + 2x + 1
```

Output:
```
#7 (generic function with 1 method)
```

```julia
sayhi3("Chewbacca")
```

Output:
```
Hi Chewbacca, it's great to see you!
```

```julia
f3(42)
```

Output:
```
1849
```

```julia
sayhi3 = "some string"
```

Output:
```
"some string"
```

```julia
sayhi2 = "some string"
```

Output:
```
Stacktrace:
 [1] top-level scope at In[16]:1
```

```julia
const my_pi = 3.14
```

Output:
```
3.14
```

```julia
my_pi = "string"
```

Output:
```
Stacktrace:
 [1] top-level scope at In[19]:1
```

## Duck-typing in Julia
*"If it quacks like a duck, it's a duck."* <br><br>
Julia functions will just work on whatever inputs make sense. <br><br>
For example, `sayhi` works on the name of this minor tv character, written as an integer...

```julia
sayhi(55595472)
```

Output:
```
Hi 55595472, it's great to see you!
```

And `f` will work on a matrix. 

```julia
A = rand(3, 3)
A
```

Output:
```
3×3 Array{Float64,2}:
 0.919815  0.504183  0.778159 
 0.77311   0.223771  0.544433 
 0.105556  0.122471  0.0476849
```

```julia
f(A)
```

Output:
```
3×3 Array{Float64,2}:
 4.15762  2.68024  3.58368
 3.4878   1.95408  2.83826
 1.40792  1.33141  1.24646
```

`f` will also work on a string like "hi" because `*` is defined for string inputs as string concatenation.

```julia
f("hi")
```

Output:
```
Stacktrace:
 [1] f(::String) at ./In[25]:2
 [2] top-level scope at In[27]:1
```

On the other hand, `f` will not work on a vector. Unlike `A^2`, which is well-defined, the meaning of `v^2` for a vector, `v`, is not a well-defined algebraic operation. 

```julia
v = rand(3)
```

Output:
```
3-element Array{Float64,1}:
 0.4200902659476402 
 0.5080408373743039 
 0.34077737444812395
```

```julia
f(v)
```

Output:
```
[2] literal_pow at ./none:0 [inlined]
 [3] f(::Array{Float64,1}) at ./In[25]:2
 [4] top-level scope at In[29]:1
```

## Mutating vs. non-mutating functions

By convention, functions ending in `!` alter their contents and functions lacking `!` do not.

For example, let's look at the difference between `sort` and `sort!`.


```julia
v = [3, 5, 2]
```

Output:
```
3-element Array{Int64,1}:
 3
 5
 2
```

```julia
sort(v)
```

Output:
```
3-element Array{Int64,1}:
 2
 3
 5
```

```julia
v
```

Output:
```
3-element Array{Int64,1}:
 3
 5
 2
```

`sort(v)` returns a sorted array that contains the same elements as `v`, but `v` is left unchanged. <br><br>

On the other hand, when we run `sort!(v)`, the contents of v are sorted within the array `v`.

```julia
sort!(v)
```

Output:
```
3-element Array{Int64,1}:
 2
 3
 5
```

```julia
v
```

Output:
```
3-element Array{Int64,1}:
 2
 3
 5
```

## Some higher order functions

### map

`map` is a "higher-order" function in Julia that *takes a function* as one of its input arguments. 
`map` then applies that function to every element of the data structure you pass it. For example, executing

```julia
map(f, [1, 2, 3])
```
will give you an output array where the function `f` has been applied to all elements of `[1, 2, 3]`
```julia
[f(1), f(2), f(3)]
```

```julia
map(f, [1, 2, 3])
```

Output:
```
3-element Array{Int64,1}:
  4
  9
 16
```

Here we've squared all the elements of the vector `[1, 2, 3]`, rather than squaring the vector `[1, 2, 3]`.

To do this, we could have passed to `map` an anonymous function rather than a named function, such as

```julia
x -> x^3
```

Output:
```
#9 (generic function with 1 method)
```

via

```julia
map(x -> x^3, [1, 2, 3])
```

Output:
```
3-element Array{Int64,1}:
  1
  8
 27
```

and now we've cubed all the elements of `[1, 2, 3]`!

```julia
[x^3 for x in [1, 2, 3]]
```

Output:
```
3-element Array{Int64,1}:
  1
  8
 27
```

```julia
v = 0:0.1:1
```

Output:
```
0.0:0.1:1.0
```

```julia
[sin(x) for x in v]
```

Output:
```
11-element Array{Float64,1}:
 0.0                
 0.09983341664682815
 0.19866933079506122
 0.29552020666133955
 0.3894183423086505 
 0.479425538604203  
 0.5646424733950354 
 0.644217687237691  
 0.7173560908995228 
 0.7833269096274834 
 0.8414709848078965
```

```julia
map(sin, v)
```

Output:
```
11-element Array{Float64,1}:
 0.0                
 0.09983341664682815
 0.19866933079506122
 0.29552020666133955
 0.3894183423086505 
 0.479425538604203  
 0.5646424733950354 
 0.644217687237691  
 0.7173560908995228 
 0.7833269096274834 
 0.8414709848078965
```

### broadcast

`broadcast` is another higher-order function like `map`. `broadcast` is a generalization of `map`, so it can do every thing `map` can do and more. The syntax for calling `broadcast` is the same as for calling `map`

```julia
broadcast(f, [1, 2, 3])
```

Output:
```
3-element Array{Int64,1}:
  4
  9
 16
```

and again, we've applied `f` (squared) to all the elements of `[1, 2, 3]` - this time by "broadcasting" `f`!

Some syntactic sugar for calling `broadcast` is to place a `.` between the name of the function you want to `broadcast` and its input arguments. For example,

```julia
broadcast(f, [1, 2, 3])
```
is the same as
```julia
f.([1, 2, 3])
```

```julia
f.([1, 2, 3])
```

Output:
```
3-element Array{Int64,1}:
  4
  9
 16
```

Notice again how different this is from calling 
```julia
f([1, 2, 3])
```
We can square every element of a vector, but we can't square a vector!

To drive home the point, let's look at the difference between

```julia
f(A)
```
and
```julia
f.(A)
```
for a matrix `A`:

```julia
A = [i + 3*j for j in 0:2, i in 1:3]
```

Output:
```
3×3 Array{Int64,2}:
 1  2  3
 4  5  6
 7  8  9
```

```julia
f(A)
```

Output:
```
3×3 Array{Int64,2}:
  33   41   49
  75   92  109
 117  143  169
```

As before we see that for a matrix, `A`,
```
f(A) = A^2 = A * A
``` 

On the other hand,

```julia
B = f.(A)
```

Output:
```
3×3 Array{Int64,2}:
  4   9   16
 25  36   49
 64  81  100
```

contains the squares of all the entries of `A`.

This dot syntax for broadcasting allows us to write relatively complex compound elementwise expressions in a way that looks natural/closer to mathematical notation. For example, we can write

```julia
A .+ 2 .* f.(A) ./ A
```

Output:
```
3×3 Array{Float64,2}:
  9.0     11.0   13.6667
 16.5     19.4   22.3333
 25.2857  28.25  31.2222
```

instead of

```julia
broadcast(x -> x + 2 * f(x) / x, A)
```

and the two will perform exactly the same.

```julia
@. A + 2 * f(A) / A
```

Output:
```
3×3 Array{Float64,2}:
  9.0     11.0   13.6667
 16.5     19.4   22.3333
 25.2857  28.25  31.2222
```

```julia
A .+ [1, 2, 3]
```

Output:
```
3×3 Array{Int64,2}:
  2   3   4
  6   7   8
 10  11  12
```

```julia
[1, 2, 3] .* [4, 5, 6]'
```

Output:
```
3×3 Array{Int64,2}:
  4   5   6
  8  10  12
 12  15  18
```

```julia
x, y = 1, 2
@assert x == y
```

Output:
```
Stacktrace:
 [1] top-level scope at In[56]:2
```

```julia
@time sleep(1)
```

Output:
```
1.030055 seconds (107.82 k allocations: 5.719 MiB)
```

```julia
@time sleep(1)
```

Output:
```
1.005414 seconds (30 allocations: 1008 bytes)
```

### Exercises

#### 6.1 
Write a function `add_one` that adds 1 to its input.

```julia
add_one(x) = x + 1
```

Output:
```
add_one (generic function with 1 method)
```

<!-- grade_id: cell-5119b9e9623c1cb7, points: 1, grade: true, solution: false -->
```julia
@assert add_one(1) == 2
```

<!-- grade_id: cell-50f83d27187a2064, points: 1, grade: true, solution: false -->
```julia
@assert add_one(11) == 12
```

#### 6.2 
Use `map` or `broadcast` to increment every element of matrix `A` by `1` and assign it to a variable `A1`.

```julia
map(add_one, A)
```

Output:
```
3×3 Array{Int64,2}:
 2  3   4
 5  6   7
 8  9  10
```

```julia
A1 = map(x -> x+1, A)
```

Output:
```
3×3 Array{Int64,2}:
 2  3   4
 5  6   7
 8  9  10
```

#### 6.3 
Use the broadcast dot syntax to increment every element of matrix `A1` by `1` and store it in variable `A2`

```julia
A2 = A1 .+ 1
```

Output:
```
3×3 Array{Int64,2}:
 3   4   5
 6   7   8
 9  10  11
```

```julia
A2 = add_one.(A1)
```

Output:
```
3×3 Array{Int64,2}:
 3   4   5
 6   7   8
 9  10  11
```

<!-- grade_id: cell-f3bd5479679a8fe1, points: 0, grade: true, solution: false -->
```julia
@assert A2 == [3 4 5; 6 7 8;9 10 11]
```

Please click on `Validate` on the top, once you are done with the exercises.

## Generator expressions

Generator expressions provide a concise way to transform and filter iterators.

```julia
ζ(k, N=1000) = sum(1/n^k for n = 1:N)
```

Output:
```
ζ (generic function with 2 methods)
```

```julia
ζ(2)
```

Output:
```
1.6439345666815615
```

```julia
@time ζ(2)
```

Output:
```
0.000007 seconds (5 allocations: 176 bytes)
1.6439345666815615
```

```julia
thunk() = all(sum(digits(i))==9 for i in 9:9:90)
```

Output:
```
thunk (generic function with 1 method)
```

```julia
@time thunk()
```

Output:
```
0.000004 seconds (14 allocations: 1.094 KiB)
true
```

```julia
words = readlines(download("http://www.mieliestronk.com/corncob_lowercase.txt"))
```

Output:
```
% Total    % Received % Xferd  Average Speed   Time    Time     Time  Current
                                 Dload  Upload   Total   Spent    Left  Speed
100  586k    0  586k    0     0   415k      0 --:--:--  0:00:01 --:--:--  415k
58110-element Array{String,1}:
 "aardvark"   
 "aardwolf"   
 "aaron"      
 "aback"      
 "abacus"     
 "abaft"      
 "abalone"    
 "abandon"    
 "abandoned"  
 "abandonment"
 "abandons"   
 "abase"      
 "abased"     
 ⋮            
 "zoological" 
 "zoologist"  
 "zoologists" 
 "zoology"    
 "zoom"       
 "zoomed"     
 "zooming"    
 "zooms"      
 "zooplankton"
 "zoos"       
 "zulu"       
 "zulus"
```

```julia
long_pals = [w for w in words if w == reverse(w) && length(w) > 4]
```

Output:
```
17-element Array{String,1}:
 "civic"  
 "deified"
 "kayak"  
 "level"  
 "madam"  
 "minim"  
 "radar"  
 "redder" 
 "refer"  
 "repaper"
 "reviver"
 "rotator"
 "rotor"  
 "sagas"  
 "sexes"  
 "shahs"  
 "tenet"
```

```julia
maximum(length(w) for w in long_pals)
```

Output:
```
7
```

```julia
maximum(length, long_pals)
```

Output:
```
7
```

### Exercise: make a table of the number of palindromes of each length

```julia
function palindrome_table(words)
    table = fill(0, 7)
    for w in words
        w == reverse(w) || continue
        table[length(w)] += 1
    end
    return table
end
```

Output:
```
palindrome_table (generic function with 1 method)
```

```julia
palindrome_table(words)
```

Output:
```
7-element Array{Int64,1}:
  0
  0
 33
  7
 12
  1
  4
```

```julia
table = fill(0, 7)

```

Output:
```
7-element Array{Int64,1}:
 0
 0
 0
 0
 0
 0
 0
```

```julia
for w in words
    w == reverse(w) || continue
    table[length(w)] += 1
end
```

```julia
table
```

Output:
```
7-element Array{Int64,1}:
  0
  0
 33
  7
 12
  1
  4
```

```julia
[length([w for w in words if w == reverse(w) && length(w) == ii]) for ii in 1:7]
```

Output:
```
7-element Array{Int64,1}:
  0
  0
 33
  7
 12
  1
  4
```

```julia
[ l => (count(w->length(w)==l && w==reverse(w), words)) for l in 1:7 ]
```

Output:
```
7-element Array{Pair{Int64,Int64},1}:
 1 => 0 
 2 => 0 
 3 => 33
 4 => 7 
 5 => 12
 6 => 1 
 7 => 4
```

```julia

```